Homework 6 : Tables

In this assignment, we will process data in tables. We will get practice with

- missing entries 
- aggregating values
- linking records
  
The questions guide you step-by-step through these approaches. 

### Rubric

Question | Points
--- | ---
Question 1.1 | 1
Question 1.2 | 1
Question 1.3 | 1
Question 2.1 | 1
Question 2.2 | 1
Question 2.3 | 1
Question 2.4 | 1
Question 3.1 | 1
Question 3.2 | 1
Question 3.3 | 1
Question 4.1 | 1
Question 4.2 | 1
Question 4.3 | 1
Question 5.1 | 1
Question 5.2 | 1
Question 5.3 | 0
Total | 15

<img src="image.PNG"  width="400" align="right"/>

We will study data from the World Bank. The World Bank is an international financial organization which provides loans to governments of developing countries. 

In Homework 3, we studied 

- literacy rate among adult females 
- gross national income per capita

Here we want to study

- gross domenstic product per capita
- population 
- life expectancy 

Note the difference between gross national income and gross domestic product.  

- Gross domestic product measures the value of domestic goods and services 
- Gross national income measures revenue generated domestically and abroad by residents and businesses.

So the gross national income tends to be higher than the gross domestic product.

### 0. Load Packages

We have been working with the

-  `numpy` package for manipulations of arrays
-  `matplotlib` package for generating charts
- `pandas` package for handling tables 

Here we will focus on operations in `pandas` including the join operation

```python
# import some packages

import numpy as np
import pandas as pd

import matplotlib.pyplot as plt

# change some settings

pd.set_option('display.max_rows', 10)
pd.set_option('display.max_columns', 8)

plt.rcParams['figure.figsize'] = (10,8)
```

<!-- grade_id: cell-873c5dce3bbf6b9f, points: 0, grade: true, solution: false -->
```python
# TEST

import sys

assert "numpy" in sys.modules and "np" in locals()
assert "pandas" in sys.modules and "pd" in locals()
assert "matplotlib" in sys.modules and "plt" in locals()
```

Note that we changed some of the default settings in `pandas` with `set_option` and in `matplotlib` with `rcParams`. 

### Question 1

Before we can understand the information about gross domestic product, we need to load and rearrange the data from `world_bank_gdp.csv`. 

#### Question 1.1

Use the `pandas` function `read_csv` to load the data from `world_bank_gdp.csv`. Store the table in the variable `gdp`.

<!-- grade_id: cell-f66819dc9874ceeb, points: 0, grade: false, solution: true -->
```python
gdp = pd.read_csv("world_bank_gdp.csv")

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-4831bca6f3d5ba65, points: 1, grade: true, solution: false -->
```python
# TEST 

assert gdp.shape == (260, 214)

```

#### Question 1.2

Use the `pandas` function `set_index` to make the column `GDP per capita` the index of the table. 

<!-- grade_id: cell-e01ce1704b50da54, points: 0, grade: false, solution: true -->
```python
column = 'GDP per capita'

# YOUR CODE HERE
#raise NotImplementedError()

gdp.set_index(column, inplace=True)
```

<!-- grade_id: cell-6f0e429ff8996bd1, points: 1, grade: true, solution: false -->
```python
# TEST 

assert gdp.index.name == "GDP per capita"

```

Now we can use `transpose` to switch the rows and columns of the table. 

```python
gdp = gdp.transpose()
```

#### Question 1.3

Remember the `pandas` function `to_datetime` from Homework 4. We want to change the year to a `datetime` object. 

Specify the format as a string. Note each entry represents a year.

<!-- grade_id: cell-b94db99ccad54446, points: 0, grade: false, solution: true -->
```python
format_string = "%Y"

# YOUR CODE HERE
#raise NotImplementedError()

gdp.index = pd.to_datetime(gdp.index, format = format_string)
```

<!-- grade_id: cell-b5b87c43f06383db, points: 1, grade: true, solution: false -->
```python
# TEST

assert all(gdp.index.year == np.arange(1800, 2013, 1))

```

### Question 2

We need to handle missing values.

#### Question 2.1

Use the `pandas` function `replace` to substitute `np.NaN` for `Missing` in the table. Remember that we need to specify a dictionary for `replace`. 

<!-- grade_id: cell-da7b2f5b92363399, points: 0, grade: false, solution: true -->
```python
dictionary_for_replacement = {"Missing":np.NaN}

# YOUR CODE HERE
#raise NotImplementedError()

gdp.replace(dictionary_for_replacement, inplace=True)
```

<!-- grade_id: cell-b7123b10dd5dc236, points: 1, grade: true, solution: false -->
```python
# TEST

assert all(~gdp.isin(["Missing"]))
```

#### Question 2.2

We want to filter the rows of the table. We will remove any year before 2000 and any year following 2010. 

Fill in the missing pieces of the logical expression for fitering `gdp`.

<!-- grade_id: cell-0452c4e6937a3983, points: 0, grade: false, solution: true -->
```python
lower_bound = 1999
upper_bound = 2011

# YOUR CODE HERE
#raise NotImplementedError()

gdp = gdp[(lower_bound < gdp.index.year) & (gdp.index.year < upper_bound)]
```

<!-- grade_id: cell-76bc5eb67fae7aba, points: 1, grade: true, solution: false -->
```python
# TEST

assert all(gdp.index.year == np.arange(2000, 2011, 1))
```

#### Question 2.3

We want to set the `pandas` function `dropna` to drop any columns containing `np.NaN`. Since we must drop columns not rows, we need to specify a value for the argument `axis`. Determine the appropriate value for `axis` to drop columns not rows.

<!-- grade_id: cell-1a3ef3b891ae286a, points: 0, grade: false, solution: true -->
```python
axis_value = 1

# YOUR CODE HERE
#raise NotImplementedError()

gdp.dropna(axis=axis_value, inplace=True)
```

<!-- grade_id: cell-51ab9eeeb261c7bc, points: 1, grade: true, solution: false -->
```python
# TEST

assert all(~gdp.isna())

```

#### Question 2.4

Use the `pandas` function `astype` to convert the data type of the entries from string to floating point number. You can specify the data type as `float64`.

<!-- grade_id: cell-484a6eb96240f6b3, points: 0, grade: false, solution: true -->
```python
floating_point_data_type = "float64"

# YOUR CODE HERE
#raise NotImplementedError()

gdp = gdp.astype(floating_point_data_type)
```

<!-- grade_id: cell-33ff8103539a6730, points: 1, grade: true, solution: false -->
```python
# TEST

assert np.issubdtype(gdp.values.dtype, np.float64)
```

### Question 3

Since we have records for 11 years, we want to aggregate the numbers.

#### Question 3.1

Use the `pandas` function `mean` to take the average of the columns.

<!-- grade_id: cell-bc5ef729d9a0e5a4, points: 0, grade: false, solution: true -->
```python
gdp_averages = gdp.mean(axis = 0)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-9793fcf2a242a5e8, points: 1, grade: true, solution: false -->
```python
# TEST

assert gdp_averages.shape == (210,)

```

#### Question 3.2

Note that `gdp_averages` is a `pandas` Series. We can use the `pandas` function `to_frame` to convert to a `pandas` DataFrame. 

```python
gdp_averages = gdp_averages.to_frame()
```

By default the name of the column is `0`. Use the `pandas` function `rename` to change from `0` to `Average GDP`. Note that we need to specify a dictionary. 

<!-- grade_id: cell-c83480050317ae49, points: 0, grade: false, solution: true -->
```python
dictionary_for_renaming = {0:"Average GDP"}

# YOUR CODE HERE
#raise NotImplementedError()

gdp_averages.rename(columns = dictionary_for_renaming, inplace = True)
```

<!-- grade_id: cell-2920dea29cee7cf1, points: 1, grade: true, solution: false -->
```python
# TEST

assert all(gdp_averages.columns == "Average GDP")
```

#### Question 3.3

Use the `pandas` function `max` to compute the maximum value in the `Average GDP` column of `gdp_averages`.

<!-- grade_id: cell-25a497c04f12b473, points: 0, grade: false, solution: true -->
```python
gdp_averages_maximum = gdp_averages.values.max(axis = 0)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-4eea5763752ef234, points: 1, grade: true, solution: false -->
```python
# TEST

assert 80000 < gdp_averages_maximum < 90000

```

Now we can standardize the values in `gdp_averages`. 

```python
gdp_averages_standardized = gdp_averages / gdp_averages_maximum
```

Having divided by `gdp_averages_maximum`, we have values in `gdp_averages_standardized` between 0 and 1.

### Question 4

Before we can understand the information about population and life expectancy, we need to load and join the data 

#### Question 4.1

Use the `pandas` function `read_csv` to load the data from `world_bank_country.csv`. Store the table in the variable `country`.

<!-- grade_id: cell-b4b45c2c870c2cf2, points: 0, grade: false, solution: true -->
```python
country = pd.read_csv("world_bank_country.csv")

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-384d4051d0f7d7fb, points: 1, grade: true, solution: false -->
```python
# TEST 

assert country.shape == (248, 17)

```

#### Question 4.2

We are interested in three columns of the table `country`.

```python
country = country[["Country (en)", "Population", "Life expectancy"]]
```

Use the `pandas` function `set_index` to make the column `Country (en)` the index of the table. 

<!-- grade_id: cell-9bbe18f220c92181, points: 0, grade: false, solution: true -->
```python
column = 'Country (en)'

# YOUR CODE HERE
#raise NotImplementedError()

country.set_index(column, inplace=True)
```

<!-- grade_id: cell-bb74c42d50e0e800, points: 1, grade: true, solution: false -->
```python
# TEST 

assert country.index.name == "Country (en)"

```

#### Question 4.3

We can the `pandas` function `merge` to join `gdp_averages_standardized` and `country` on their indices. Remember to set

- `left_index=True`
- `right_index=True`

to indicate that we join on the indices.

<!-- grade_id: cell-69ee60f80a25905e, points: 0, grade: false, solution: true -->
```python
gdp_country_merged = pd.merge(gdp_averages_standardized,country, left_index=True, right_index=True)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-c875172751fbdf33, points: 1, grade: true, solution: false -->
```python
assert gdp_country_merged.shape == (194, 3)

```

### Question 5

Having processed the data in `world_bank_country.csv` and `world_bank_gdp.csv`, we can generate a scatter-plot showing average gross domestic product per capita, life expectancy and population.

#### Question 5.1 

Use the `pandas` function `sort_values` to sort `gdp_country_merged` in descending order by `Average GDP`.

<!-- grade_id: cell-65ca7feec2f0b81e, points: 0, grade: false, solution: true -->
```python
order = False

# YOUR CODE HERE
#raise NotImplementedError()

gdp_country_merged.sort_values('Average GDP', ascending=order, inplace=True)
```

<!-- grade_id: cell-b3c6a320d9d65ebd, points: 1, grade: true, solution: false -->
```python
# TEST 

assert gdp_country_merged['Average GDP'].is_monotonic_decreasing
```

#### Question 5.2

Since some countries have a large population and some countries have a small population, we will use the logarithm function from Homework 3.

Use the `numpy` function `log` to transform the entries in the `Population` column.

<!-- grade_id: cell-8bcc4c0ef5b0d61e, points: 0, grade: false, solution: true -->
```python
logarithm_function = np.log

# YOUR CODE HERE
#raise NotImplementedError()

gdp_country_merged["Log Population"] = logarithm_function(gdp_country_merged["Population"])
```

<!-- grade_id: cell-d1d6c4a77b4055bc, points: 1, grade: true, solution: false -->
```python
# TEST 

assert "Log Population" in gdp_country_merged.columns

```

#### Question 5.3

Now we can generate a scatter-plot showing the information in `gdp_country_merged`.

- Horizontal coordinate is `Log Population`
- Vertical coordinate is `Life Expectancy`
- Size is determined by `Average GDP`

```python
plt.scatter(gdp_country_merged["Log Population"], 
            gdp_country_merged["Life expectancy"], 
            s = gdp_country_merged['Average GDP']*200, label="Average GDP");

plt.ylim([45, 95])

plt.ylabel("Life expectancy")
plt.xlabel("Log of Popualtion")
plt.title("World Bank Data")

plt.legend();
```