 Homework 2 : Classes


In this assignment, we will learn to combine containers and functions into classes. We will explore ways to make connections between classes through inheritance. 

The questions guide you step-by-step through these approaches. Please post to Slack with any questions under the #homework-2 channel. 


### Rubric

Question | Points
--- | ---
Question 1.1 | 1
Question 1.2 | 1
Question 1.3 | 1
Question 2 | 1
Question 3 | 1
Question 4 | 0
Question 5.1 | 1
Question 5.2 | 1
Question 5.3 | 1
Total | 8

We discussed research of Edmond Halley in Lecture 1. Remember that Edmond Halley used demographics data from Poland to estimate life expectancy at different ages. Using these estimates for life expectancy, Edmond Hally could determine a price for annuities in England. Remember that annuities are contracts in the insurance industry. 

Here we will study the price of bonds. Bonds are contracts in the finance industry. We can understand the price of bonds in terms of schedule, payments and discounts. Following the jargon of the finance industry, we will refer to a schedule of payments as a cash-flow. 

##### Schedule 

Since bonds are contracts, bonds have a starting time and an ending time. We will study contracts with starting time $0$ years and ending time $T$ years. So the contract ends $T$ years from the present time. The schedule consists of years 

> $1,2,\ldots,T-1,T$ years.

For each year in the schedule, the bond has a payment.

##### Payment

The payment depends on the principal of the bond and on the coupon of the bond. Let $P$ denote the principal of the bond. Let $c$ denote the coupon of the bond. For each year in the schedule, the bond pays 

> $c \cdot P$ dollars 

Additionally the bond pays the principal at the ending time $T$. The contract behind the bond depends on the schedule and payments. However the price of the bond requires more information than the schedule and payments. The missing components come from financial markets. 

##### Discount 

The missing components of the price of bonds are discount factors. If we study a payment at a future time, then we need to discount the payment by some factor to determine the value at the present time. 

We can understand discounting by analogy with inflation. If a product costs $\$1$ at the present time, it might cost $\$10$ at time $T$ years. So a payment of $\$10$ at year $T$ would need to be discounted by a factor of $0.1$ to obtain a payment of \$1 at the present time.  

We have two ways to specify a discount factor for time $T$ years using a discount rate $r$

> - Simple Discounting
>   * $\left(\frac{1}{1 + r}\right)^T$
> - Continuous Discounting
>   * $e^{-r T}$ where $e = 2.7182\ldots$

Let $D_t$ denote the discount factor at time $t$ years. 

By combining schedule, payment, and discounts, we can determine the price of the bond $$\text{Price} =  P \cdot D_T + \sum_{t = 1}^{T}\;  c \cdot  P \cdot D_t \ . $$ Please check [here](https://en.wikipedia.org/wiki/Summation) for information about the notation.  For example, take 

 - $P = 100$ dollars
 - $T = 3$ years 
 - $c = 0.05$ coupon 
 - $r_1 = 0.04$, $r_2 = 0.05$, $r_3 = 0.06$ rates
 
For simple discounting, we would have 

$$97.5 = 100 \cdot \left(\frac{1}{1 + 0.06}\right)^3 + 0.05 \cdot 100 \cdot \left(\frac{1}{1 + 0.04}\right)^1 + 0.05 \cdot 100 \cdot \left(\frac{1}{1 + 0.05}\right)^2 + 0.05 \cdot 100 \cdot \left(\frac{1}{1 + 0.06}\right)^3 $$  

For continuous discounting, we would have 

$$97.0 = 100 \cdot e^{-0.06 \cdot 3}  + 0.05 \cdot 100 \cdot e^{-0.04 \cdot 1} + 0.05 \cdot 100 \cdot e^{-0.05 \cdot 2} + 0.05 \cdot 100 \cdot e^{-0.06 \cdot 3} $$  

### Question 1. Function for Calculating Bond Prices  

We want to write functions to compute the price of bonds 

- a function that computes the price from cash-flows and discount factors 
- a function to generate discount factors from rates

Combining these functions, we will be able to determine the price of a bond.

#### Question 1.1

We want to write a function called `price_product`. The inputs are 
 
 - ```C``` a list of cash-flows
 - ```D``` a list of discount factors
 
We assume `C` and `D` have the same length. The output is a number. The number is the sum of the products 

- `C[i] * D[i]`
 
Here `i` ranges from 0 to the length `C` or equivalently the length of `D`. Complete the following code to implement the function.

<!-- grade_id: cell-584e947774ffa149, points: 0, grade: false, solution: true -->
```python
def price_product(C, D):
    output = 0
    x=len(C)
    
    for i in range(0,x):
        output = C[i]*D[i] + output
    
    return output

```

<!-- grade_id: cell-3173068f98658409, points: 1, grade: true, solution: false -->
```python
# TEST 

C = [5.0, 5.0, 105.0]
D = [0.96, 0.90, 0.83]

assert type(price_product(C,D)) == float

```

#### Question 1.2

We want to write a function called ```generate_simple_discount``` with inputs 
 
 - ```r``` rates

Here the rates `r` should be a list. The output is a list. The list contains the discount factors 

$$\left(\frac{1}{1 + r_1}\right)^1,\left(\frac{1}{1 + r_2}\right)^2,\ldots,\left(\frac{1}{1 + r_{T-1}}\right)^{T-1}, \left(\frac{1}{1 + r_T}\right)^T $$ 

Note that $T$ should be the length of the list with $r_1$ the entry with index 0, $r_2$ the entry with index 1,...

<!-- grade_id: cell-e7a71350541efc32, points: 0, grade: false, solution: true -->
```python
def generate_simple_discount(r):
    output = []

    for i in range(0,len(r)):
        t=i+1
        output.append(pow(1/(1+r[i]),t))
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    return output
```

<!-- grade_id: cell-0df2eadcb6c0fba4, points: 1, grade: true, solution: false -->
```python
# TEST 

r = [0.04, 0.05, 0.06]
discount_factors = generate_simple_discount(r)

assert abs(discount_factors[0] - 0.96) < 1

```

#### Question 1.3 

We want to compute the price of a bond 
 
 - $P = 1000$ dollars
 - $T = 3$ years 
 - $c = 0.03$ coupon 
 - $r_1 = 0.07$, $r_2 = 0.08$, $r_3 = 0.09$ rates 

We need to determine the cash-flows.

```python
P = 1000 
c = 0.03
T = 3
r = [0.07, 0.08, 0.09]

C = T * [c * P]
C[-1] = C[-1] + P
```

Note that we have to add `P` to the last entry of the list because the bond pays the principal at the ending time.
 
Use `price_product` from Question 1.1 and `generate_simple_discount` from Question 1.2 to calculate the price of the bond. Note you answer should equal

$$1000 \cdot \left(\frac{1}{1 + 0.09}\right)^3 + 0.03 \cdot 1000 \cdot \left(\frac{1}{1 + 0.07}\right)^1 + 0.03 \cdot 1000 \cdot \left(\frac{1}{1 + 0.08}\right)^2 + 0.03 \cdot 1000 \cdot \left(\frac{1}{1 + 0.09}\right)^3 $$  

<!-- grade_id: cell-edfb3298b3ff9cdf, points: 0, grade: false, solution: true -->
```python
discount_factors1 = generate_simple_discount(r)
q1_3 = price_product(C,discount_factors1)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-840c0f67a57d7f3b, points: 1, grade: true, solution: false -->
```python
# TEST 


assert 750 < q1_3 < 1000

```

### Question 2 : Class for Schedule and Payment

We want a class to store the schedule and payment of a bond.

```python
class Bond(object):
    def __init__(self, P, T, c):
        self.cash_flows = self.generate_cash_flows(P, T, c)
    
    def generate_cash_flows(self, P, T, c):
        output = T * [c * P]
        output[-1] = output[-1] + P
        return output
```

Here we have a class called `Bond`. The class has two methods 

- `__init__` makes an object from the class 
- `generate_cash_flows` determines the cash-flows

Rememeber that cash-flows means the schedule of payments. The function `__init__` calls the function `generate_cash_flows` with the values of the principal `P`, ending time `T` and coupon `c` to determine the payments 

> `[c * P, c * P,...,c * P, c * P + P]` 

Here the list has `T` entries.

Create an object from the `Bond` class  with 

 - $P = 1000$ dollars
 - $T = 3$ years 
 - $c = 0.03$ coupon 

Call the object `bond_object`.

<!-- grade_id: cell-6d341117143a840e, points: 0, grade: false, solution: true -->
```python
P = 1000
T = 3
c = 0.03

bond_object = Bond(P,T,c)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-cab182f8645252a7, points: 1, grade: true, solution: false -->
```python
# TEST 

assert len(bond_object.cash_flows) == 3

```

### Question 3 : Class for Discount

We would have to use market conditions to determine discount factors. Practitioners tend to use discount factors determined from government bonds like US Treasury bonds with ending time 1 year, 5 year or 10 year.

<img src="image_2.png">

Since we have a class to store the schedule and payment of the bond, we need a class to determine the discount factors. 

```python
class Discounter(object):
    def __init__(self, rates):
        self.rates = rates 
        
    def get_discount_factors(self):
        raise(NotImplementedError)
```

Here we have a class called `Discounter`. The class has two methods 

- `__init__` makes an object from the class 
- `get_discount_factors` does not have an implementation

We can determine the discount factors from the rates. However, we need to specify simple discounting or continuous discounting. We can use a child class to implement `get_discount_factors`.

```python
class DiscounterContinuous(Discounter):
    def __init__(self, rates):
         super().__init__(rates)
        
    def get_discount_factors(self):        
        constant = 2.7182
        
        discount_factors = []
        for i,r in enumerate(self.rates): 
            discount_factors.append(constant**(-r * (i+1)))
        
        return discount_factors
```

Here we have a class called `DiscounterContinuous`. The class has two methods 

- `__init__` makes an object from the class by passing `rates` to the parent class 
- `get_discount_factors` determines discount factors from `rates`

Use `generate_simple_discount` from Question 1.2 to complete the code for a class `DiscounterSimple`.

<!-- grade_id: cell-08f7c5b916ea3087, points: 0, grade: false, solution: true -->
```python
class DiscounterSimple(Discounter):
    def __init__(self, rates):
         super().__init__(rates)
        
    def get_discount_factors(self):            
        discount_factors = generate_simple_discount(r)
        
        # YOUR CODE HERE
        #raise NotImplementedError()
        
        return discount_factors
```

<!-- grade_id: cell-a0f2e989ff993b2f, points: 1, grade: true, solution: false -->
```python
# TEST 

r = [0.04, 0.05, 0.06]
discounter_simple_object = DiscounterSimple(r)
discount_factors = discounter_simple_object.get_discount_factors()

assert abs(discount_factors[0] - 0.96) < 1

```

### Question 4 : Class for Pricing Products

We want to have a class that prices bonds. However, we want to make the class useful for calculating the price of financial products that resemble bonds. 

Remember that the schedule and payments in the contract for the bond determine the cash-flows. We want the class that prices bonds to handle any financial product with cash-flows. 

```python
class Pricer(object):
    def __init__(self, Discounter):
        self.discounts = Discounter
        
    def price(self, Product):
        cash_flows = Product.get_cash_flows() 
        discount_factors = self.discounts.get_discount_factors()

        return price_product(cash_flows, discount_factors)
```

Here we have a class called `Pricer`. The class has two methods 

- `__init__` makes an object from the class 
- `price` determines the price of a `Product` object

Since we need discount factors to determine the price of the financial product from the cash-flows, we need to use the `Discounter` class. If we make an object of the `Pricer` class, then we need to pass an object from the `Discounter` class to `__init__`. 

Note that we pass an object from the `Product` class to `price`. We need to obtain the cash-flows from the object. With the cash-flows and discount factors, we can compute the price of the financial product with `price_product` function from Question 1.1.

### Question 5 : Class for Pricing Bonds

We want to modify the `Bond` class from Question 2. We will have the `Bond` class inherit from a `Product` class. We should be able to price any object in the `Product` class. So we need to have a `get_cash_flows` from Question 4.

```python
class Product(object):
    def __init__(self):
        self.cash_flows = []
        
    def get_cash_flows(self):
        return self.cash_flows
        
    def generate_price(self, Pricer):
        return Pricer.price(self)

class Bond(Product):
    def __init__(self, P, T, c):
        super().__init__()
        self.cash_flows = self.generate_cash_flows(P, T, c)
        
    def generate_cash_flows(self, P, T, c):
        output = T * [c * P]
        output[-1] = output[-1] + P
        return output        
```

Here we have a class called `Product`. The class has three methods 

- `__init__` makes an object from the class
- `get_cash_flows` accesses the list `cash_flows` stored in the object
- `generate_price` determines price with the help of an object from the `Pricer` class

The class `Product` is the parent class to the child class `Bond`. We have not changed the code for the child class `Bond` from Question 2 with the exception of calling `__init__` from the parent class `Product` to make an object. 

We can visualize the connections between classes with a diagram. See Instructional Video Module 2 for more information about the diagram. 

<img src="image_1.PNG">

We want to determine the price of an object from the `Bond` class with 
 
 - $P = 100$ dollars
 - $T = 3$ years 
 - $c = 0.05$ coupon 
 - $r_1 = 0.04$, $r_2 = 0.05$, $r_3 = 0.06$ rates

We will create an object from the `Pricer` class and an object from the `Bond` class to determine the prices.

#### Question 5.1 

Create an object  from the `DiscounterSimple` class with 

 - $r_1 = 0.04$, $r_2 = 0.05$, $r_3 = 0.06$ rates

Call the object `discounter_object`.

<!-- grade_id: cell-ea0a5dd857526ab3, points: 0, grade: false, solution: true -->
```python
r = [0.04, 0.05, 0.06]

discounter_object = DiscounterSimple(r)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-f0a7f3f99fa2e118, points: 1, grade: true, solution: false -->
```python
# TEST 

assert discounter_object.rates == r
```

#### Question 5.2 

Create an object  from the `Pricer` class with `discounter_object`. Call the object `pricer_object`.

<!-- grade_id: cell-e10fd7818db33b12, points: 0, grade: false, solution: true -->
```python
pricer_object = Pricer(discounter_object)

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-7b5e1636bb496526, points: 1, grade: true, solution: false -->
```python
# TEST 

assert pricer_object.discounts == discounter_object
```

#### Question 5.3

Create an object from the `Bond` class with 

 - $P = 100$ dollars
 - $T = 3$ years 
 - $c = 0.05$ coupon 

Call the object `bond_object`.

<!-- grade_id: cell-1877b77278ddd8eb, points: 0, grade: false, solution: true -->
```python
P = 100 
T = 3
c = 0.05

bond_object = Bond(P,T,c)

# YOUR CODE HERE
#raise NotImplementedError()
```

Now we can generate the price.

```python
price = bond_object.generate_price(pricer_object)
```

<!-- grade_id: cell-6896ae0809e40b81, points: 1, grade: true, solution: false -->
```python
# TEST 

assert 90 < price < 100

```